<!-- grade_id: set_a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
a = 1
```

<!-- grade_id: foo, points: 1.0, grade: true, solution: false -->
```python
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: bar, points: 1.0, grade: true, solution: false -->
```python
assert a == 1
```

<!-- grade_id: baz, points: 2.0, grade: true, solution: true -->
An answer!

<!-- grade_id: quux, points: 3.0, grade: true, solution: true -->
```python
# YOUR CODE HERE
b = 2
```

```python
print("Don't change this cell!")
```

Output:
```
Don't change this cell!
```

This cell shouldn't be changed.