<!-- grade_id: set_a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-1-a3eb14613b08> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: foo, points: 1.0, grade: true, solution: false -->
```python
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: bar, points: 1.0, grade: true, solution: false -->
```python
assert a == 1
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-0e0a66a72109> in <module>()
----> 1 assert a == 1

NameError: name 'a' is not defined
```

<!-- grade_id: baz, points: 2.0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: quux, points: 3.0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-4-a3eb14613b08> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
#print("Don't change this cell!")
```

This cell shouldn't 
be changed.